An important role for triglyceride in regulating spermatogenesis

Author Response

The following is the authors’ response to the previous reviews.

PUBLIC REVIEWS

Reviewer #1 (Public Review):

In this study, the authors investigate the role of triglycerides in spermatogenesis. This work is based on their previous study (PMID: 31961851) on triglyceride sex differences in which they showed that somatic testicular cells play a role in whole body triglyceride homeostasis. In the current study, they show that lipid droplets (LDs) are significantly higher in the stem and progenitor cell (pre-meiotic) zone of the adult testis than in the meiotic spermatocyte stages. The distribution of LDs anti-correlates with the expression of the triglyceride lipase Brummer (Bmm), which has higher expression in spermatocytes than early germline stages. Analysis of a bmm mutant (bmm[1]) - a P-element insertion that is likely a hypomorphic - and its revertant (bmm[rev]) as a control shows that bmm acts autonomously in the germline to regulate LDs. In particular, the number of LDs is significantly higher in spermatocytes from bmm[1] mutants than from bmm[rev] controls. Testes from males with global loss of bmm (bmm[1]) are shorter than controls and have fewer differentiated spermatids. The zone of bam expression, typically close to the niche/hub in WT, is now many cell diameters away from the hub in bmm[1] mutants. There is an increase in the number of GSCs in bmm[1] homozygotes, but this phenotype is probably due to the enlarged hub. However, clonal analyses of GSCs lacking bmm indicate that a greater percentage of the GSC pool is composed of bmm[1]-mutant clones than of bmm[rev]-clones. This suggests that loss of bmm could impart a competitive advantage to GSCs, but this is not explored in greater detail. Despite the increase in number of GSCs that are bmm[1]-mutant clones, there is a significant reduction in the number of bmm[1]-mutant spermatocyte and post-meiotic clones. This suggests that fewer bmm[1]mutant germ cells differentiate than controls. To gain insights into triglyceride homeostasis in the absence of bmm, they perform mass spec-based lipidomic profiling. Analyses of these data support their model that triglycerides are the class of lipid most affected by loss of bmm, supporting their model that excess triglycerides are the cause of spermatogenetic defects in bmm[1]. Consistent with their model, a double mutant of bmm[1] and a diacylglycerol Oacyltransferase 1 called midway (mdy) reverts the bmm-mutant germline phenotypes.

There are numerous strengths of this paper. First, the authors report rigorous measurements and statistical analyses throughout the study. Second, the authors utilize robust genetic analyses with loss-of-function mutants and lineage-specific knockdown. Third, they demonstrate the appropriate use of controls and markers. Fourth, they show rigorous lipidomic profiling. Lastly, their conclusions are appropriate for the results. In other words, they don't over-state the results. Overall, the rigorously quantified results support the major aim that appropriate regulation of triglycerides are needed in a germline cell-autonomous manner for spermatogenesis.

This paper should have a positive impact on the field. First and foremost, there is limited knowledge about the role of lipid metabolism in spermatogenesis. The lipidomic data will be useful to researchers in the field who study various lipid species. Going forward, it will be very interesting to determine what triglycerides regulate in germline biology. In other words, what functions/pathways/processes in germ cells are negatively impacted by elevated triglycerides. And as the authors point out in the discussion, it will be important to determine what regulates bmm expression such that bmm is higher in later stages of germline differentiation.

We thank the Reviewer for their positive assessment of our revised manuscript!

Reviewer #2 (Public Review):

Summary:

Here, the authors show that neutral lipids play a role in spermatogenesis. Neutral lipids are components of lipid droplets, which are known to maintain lipid homeostasis, and to be involved in non-gonadal differentiation, survival, and energy. Lipid droplets are present in the testis in mice and Drosophila, but not much is known about the role of lipid droplets during spermatogenesis. The authors show that lipid droplets are present in early differentiating germ cells, and absent in spermatocytes. They further show a cell autonomous role for the lipase brummer in regulating lipid droplets and, in turn, spermatogenesis in the Drosophila testis. The data presented show that a relationship between lipid metabolism and spermatogenesis is congruous in mammals and flies, supporting Drosophila spermatogenesis as an effective model to uncover the role lipid droplets play in the testis.

Strengths and weaknesses:

The authors do a commendably thorough characterization of where lipid droplets are detected in normal testes: located in young somatic cells, and early differentiating germ cells. They use multiple control backgrounds in their analysis, including w[1118], Canton S, and Oregon R, which adds rigor to their interpretations. The authors employ markers that identify which lipid droplets are in somatic cells, and which are in germ cells. The authors use these markers to present measured distances of somatic and germ cell-derived lipid droplets from the hub. Because they can also measure the distance of somatic and germ cells with age-specific markers from the hub, these results allow the authors to correlate position of lipid droplets with the age of cells in which they are present. This analysis is clearly shown and well quantified.

The quantification of lipid droplet distance from the hub is applied well in comparing brummer mutant testes to wild type controls. The authors measure the number of lipid droplets of specific diameters, and the spatial distribution of lipid droplets as a function of distance from the hub. These measurements quantitatively support their findings that lipid droplets are present in an expanded population of cells further from the hub in brummer mutants. The authors further quantify lipid droplets in germline clones of specified ages; the quantitative analysis here is displayed clearly and supports a cell autonomous role for brummer in regulating lipid droplets in spermatocytes.

Data examining testis size and number of spermatids in brummer mutants clearly indicates the importance of regulating lipid droplets to spermatogenesis. The authors show beautiful images supported by rigorous quantification supporting their findings that brummer mutants have both smaller testes with fewer spermatids at both 29 and 25C. There is also significant data supporting defects in testis size, but not spermatid number, in 14-day-old brummer mutant animals compared to controls. Their analysis clearly shows an expanded region beyond the testis apex that includes younger germ cells, supporting a role for lipid droplets influencing germ cell differentiation during spermatogenesis.

The authors present a series of data exploring a cell autonomous role for brummer in the germline, including clonal analysis and tissue specific manipulations. The clonal data indicating increased lipid droplets in spermatocyte clones, and a higher proportion of brummer mutant GSCs at the hub are convincing and supported by quantitation. The authors also show a tissue specific rescue of the brummer testis size phenotype by knocking down mdy specifically in germ cells, which is also supported by statistically significant quantitation. The authors present data examining the number of spermatocyte and post-meiotic clones 14 days after clonal induction. Their finding is significant with a p-value of 0.0496, which they acknowledge is less robust than their other data reported in this study, and could be a result of a low sample size. They indicate that future studies might validate these results with additional samples.

The authors do a beautiful job of validating where they detect brummer-GFP by presenting their own pseudotime analysis of publicly available single cell RNA sequencing data. Their data is presented very clearly, and supports expression of brummer in older somatic and germline cells of the age when lipid droplets are normally not detected. The authors also present a thorough lipidomic analysis of animals lacking brummer to identify triglycerides as an important lipid droplet component regulating spermatogenesis.

Impact:

The authors present data supporting the broad significance of their findings across phyla. This data represents a key strength of this manuscript. The authors show that loss of a conserved triglyceride lipase impacts testis development and spermatogenesis, and that these impacts can be rescued by supplementing diet with medium-chain triglycerides. The authors point out that these findings represent a biological similarity between Drosophila and mice, supporting the relevance of the Drosophila testis as a model for understanding the role of lipid droplets in spermatogenesis. The connection buttresses the relevance of these findings and this model to a broad scientific community.

We thank the Reviewer for their positive assessment of our revised paper!

RECOMMENDATIONS FOR THE AUTHORS

Reviewer #2 (Recommendations For The Authors):

The authors addressed most of my recommendations in a way that is satisfactory to me. I would like a bit more information added to the methods section about how hub area was quantified. For example, did the authors measure area within a defined region in a single Z plane (perhaps the Z plane at the center of the hub, or the Z plane with the largest area)? Alternatively, did they authors measure area in a more 3 dimensional way, i.e. volume. Adding this information to the methods would satisfy all of my previous recommendations.

We thank the Reviewer for pointing out that this information was not clear in the revised manuscript. We changed the methods section to clarify our methods as follows:

“The hub was identified as the FasIII-positive area of the testis. Hub size was estimated by measuring the FasIII-positive area in a Z-projected image of the hub in each testis. Zprojections were made using the ‘sum slices’ function in Fiji.”

Changes from the earlier version to the revised version

Abstract

Abstract

Drosophila is a powerful model to study how lipids affect spermatogenesis. Yet, the contribution of neutral lipids, a major lipid group which resides in organelles called lipid droplets (LD), to sperm development is largely unknown. Emerging evidence suggests LD are present in the testis and that loss of neutral lipid- and LD-associated genes causes subfertility; however, key regulators of testis neutral lipids and LD remain unclear. Here, we show LD are present in early-stage somatic and germline cells within the Drosophila testis. We identified a role for triglyceride lipase brummer (bmm) in regulating testis LD, and found that whole-body loss of bmm leads to defects in sperm development. Importantly, these represent cell-autonomous roles for bmm in regulating testis LD and spermatogenesis. Because lipidomic analysis of bmm mutants revealed excess triglyceride accumulation, and spermatogenic defects in bmm mutants were rescued by genetically blocking triglyceride synthesis, our data suggest that bmm-mediated regulation of triglyceride influences sperm development. This identifies triglyceride as an important neutral lipid that contributes to Drosophila sperm development, and reveals a key role for bmm in regulating testis triglyceride levels during spermatogenesis.

Introduction

Lipids play an essential role in regulating spermatogenesis across animals [1–4]. Studies in Drosophila have illuminated key roles for multiple lipid species in regulating sperm development [5–7]. For example, phosphatidylinositol and its phosphorylated derivatives participate in diverse aspects of Drosophila spermatogenesis including meiotic cytokinesis [1,8–11], somatic cell differentiation [12], germline and somatic cell polarity maintenance [13–16], and germline stem cell (GSC) maintenance and proliferation [17]. Membrane lipids also influence sperm development [18,19], whereas fatty acids play a role in processes such as meiotic cytokinesis [20] and sperm individualization [21,22]. While these studies suggest key roles for membrane lipids and fatty acids during Drosophila spermatogenesis, some of which are conserved in mammals [23–25], much less is known about how neutral lipids contribute to spermatogenesis.

Neutral lipids are a major lipid group that includes triglyceride and cholesterol ester, and reside within specialized organelles called lipid droplets (LD) [26]. LD are found in diverse cell types (e.g. adipocytes, muscle, liver, glia, neurons) [27,28,26], and play key roles in maintaining cellular lipid homeostasis. In nongonadal cell types, correct regulation of LD contributes to cellular energy production [29–31], sequestration and redistribution of lipid precursors [32–36], and regulation of lipid toxicity [37–39]. The importance of LD to normal cellular function in nongonadal cell types is shown by the fact that dysregulation of LD causes defects in cell differentiation, survival, and energy production [26,37,40,41]. In the testis, much less is known about the regulation and function of neutral lipids and LD, and how this regulation affects sperm development.

Multiple lines of evidence suggest a potential role for neutral lipids and LD during spermatogenesis. First, genes that encode proteins associated with neutral lipid metabolism and LD are expressed in the testis across multiple species [42–44]. Second, testis LD have been identified in mammals and flies under both normal physiological conditions [27,44–48] and after mitochondrial stress [49]. Third, loss of genes associated with neutral lipid metabolism and LD cause subfertility phenotypes in both flies and mammals [27,50–52]. While studies suggest that mammalian testis LD contribute to steroidogenesis [53,54], the spatial, temporal, and cell-type specific requirements for neutral lipids and LD in the testis have not been explored in detail in any animal. It remains similarly unclear which genes are responsible for regulating neutral lipids and LD during spermatogenesis.

Earlier version: To address these knowledge gaps, we used Drosophila to investigate the regulation and function of neutral lipids and LD during sperm development. Our detailed analysis of spermatogenesis under normal physiological conditions revealed the presence of LD in early-stage somatic and germline cells in the testis. We identified triglyceride lipase brummer (bmm) as a regulator of testis LD, and showed that this represents a cell-autonomous role for bmm in the germline. Importantly, we found that bmm-mediated regulation of testis LD was significant for spermatogenesis, as both whole-body and cell-autonomous loss of bmm caused defects in sperm development.

Revised version: To address these knowledge gaps, we used Drosophila to investigate the regulation and function of neutral lipids and LD during sperm development. Our detailed analysis of spermatogenesis under normal physiological conditions revealed the presence of LD in early-stage somatic and germline cells in the testis. We identified triglyceride lipase brummer (bmm) as a regulator of testis LD, and showed that this represents a cell-autonomous role for bmm in the germline. Importantly, we found that bmm-mediated regulation of testis LD was significant for spermatogenesis, as both whole-body and cell-autonomous loss of bmm caused defects in sperm development. Given that our lipidomic analysis revealed an excess accumulation of triglyceride in animals lacking bmm, and that genetically blocking triglyceride synthesis rescued many spermatogenic defects associated with bmm loss, our data suggests that bmm-mediated regulation of triglyceride is important for normal Drosophila sperm development. This reveals previously unrecognized roles for neutral lipids such as triglyceride in regulating spermatogenesis, and for bmm in regulating sperm development under normal physiological conditions. Together, these findings advance knowledge of the regulation and function of neutral lipids during spermatogenesis.

Removed in the revised version: Given that our lipidomic analysis revealed an excess accumulation of triglyceride in animals lacking bmm, and that genetically blocking triglyceride synthesis rescued many spermatogenic defects associated with bmm loss, our data suggests that bmm-mediated regulation of triglyceride is important for normal Drosophila sperm development. This reveals previously unrecognized roles for neutral lipids such as triglyceride in regulating spermatogenesis, and for bmm in regulating sperm development under normal physiological conditions. Together, these findings advance knowledge of the regulation and function of neutral lipids during spermatogenesis.

Results

Lipid droplets are present in early-stage somatic and germline cells

Earlier version: We previously reported the presence of small LD (<1 μm) at the apical tip of the testis [27], a finding we reproduced in w1118males using neutral lipid stain BODIPY (Figure 1A). These LD were present in the region that contains stem cells, early-stage somatic cells, and germline cells (Figure 1A-A’, arrows). LD were also present in the hub, an organizing center and stem cell niche in the Drosophila testis (Figure 1A’’-A’’’, arrows) [55], but largely absent within the area occupied by spermatocytes (Figure 1A and A’, arrowheads). This LD distribution was reproduced in two independent genetic backgrounds and at two additional ages (Figure 1B, 1C). While LD may contain multiple neutral lipid species [56], cholesterol-binding fluorescent polyene antibiotic filipin III did not detect cholesterol within testis LD (Figure S1A), suggesting triglyceride is the main neutral lipid in Drosophila testis LD.

Revised version: We previously reported the presence of small LD (<1 μm) at the apical tip of the testis [27], a finding we reproduced in w1118 males using neutral lipid stain BODIPY (Figure 1A). These LD were present in the region that contains stem cells, early-stage somatic cells, and germline cells (Figure 1A-A’, arrows). LD were also present in the hub, an organizing center and stem cell niche in the Drosophila testis (Figure 1A’’-A’’’, arrows) [55], but largely absent within the area occupied by spermatocytes (Figure 1A and A’, arrowheads). This LD distribution was reproduced in two independent genetic backgrounds and at two additional ages (Figure 1B, 1C). While LD may contain multiple neutral lipid species [56], cholesterol-binding fluorescent polyene antibiotic filipin III did not detect cholesterol within testis LD (Figure S1A), suggesting triglyceride is the main neutral lipid in Drosophila testis LD.

Lipid droplets are present in early-stage somatic and germ cells.

Earlier version: (A) Testis lipid droplets (LD) in w1118 animals visualized with neutral lipid dye BODIPY. (A,A’) Scale bar=50 μm; (A’’,A’’’) scale bar=15 μm. Asterisk indicates hub in all images. Arrows point to LD; arrowheads point to spermatocytes in A,B. Spermatocytes were identified as described in methods section. (B) Testis LD visualized with BODIPY in newly-eclosed males from two wild-type genotypes. Scale bars: main image=50 μm; inset image=10 μm. (C) Testis LD from w1118animals at different times post-eclosion. Scale bars=50 μm. (D) Testis LD visualized with LipidTox Red in animals with somatic cell overexpression of GFP-LD (Tj-GAL4>UAS-GFP-LD). GFP- and LipidTox Red-positive punctae are somatic LD (D–D’’ arrows); LipidTox punctae without GFP indicate germline LD (D–D’’ arrowheads). Scale bars=10 μm. (E) Histogram showing the spatial distribution of somatic cell LD; error bars represent standard error of the mean (SEM). (F) Cumulative frequency distributions of somatic LD (blue line, data reproduced from E), zfh-1-positive somatic cells (zfh-1+ cells, orange line), and Eya-positive somatic cells (Eya+ cells, grey line). (G) Testis LD visualized with LipidTox Red in males with germline overexpression of GFP-LD (nos-GAL4>UAS-GFP-LD). GFP- and LipidTox Red-positive punctae indicate germline LD (arrows); LipidTox punctae without GFP indicate non-germline LD (arrowheads). Scale bars=10 μm. (H) Histogram representing the spatial distribution of LD within the germline; error bars represent SEM. (I) Histogram representing the spatial distribution of LD and GFP fluorescence (green line) (arbitrary units, a.u.) in a representative testis of a bam-GFP animal (panel J). (J) Testis LD in a bam-GFP animal; arrows point to LD and arrowheads point to spermatocytes. Scale bar=50 μm. See also Supplemental Figure 1.

Revised version: (A) Testis lipid droplets (LD) in w1118 animals visualized with neutral lipid dye BODIPY. (A,A’) Scale bar=50 μm; (A’’,A’’’) scale bar=15 μm. Asterisk indicates hub in all images. Arrows point to LD; arrowheads point to spermatocytes in A,B. Spermatocytes were identified as described in methods section. (B) Testis LD visualized with BODIPY in newly-eclosed males from two wild-type genotypes. Scale bars: main image=50 μm; inset image=10 μm. (C) Testis LD from w1118 animals at different times post-eclosion. Scale bars=50 μm. (D) Testis LD visualized with LipidTox Red in animals with somatic cell overexpression of GFP-LD (Tj-GAL4>UAS-GFP-LD). GFP- and LipidTox Red-positive punctae are somatic LD (D–D’’ arrows); LipidTox punctae without GFP indicate germline LD (D–D’’ arrowheads). Scale bars=10 μm. (E) Histogram showing the spatial distribution of somatic cell LD; error bars represent standard error of the mean (SEM). (F) Cumulative frequency distributions of somatic LD (blue line, data reproduced from E), zfh-1-positive somatic cells (zfh-1+ cells, orange line), and Eya-positive somatic cells (Eya+ cells, grey line). (G) Testis LD visualized with LipidTox Red in males with germline overexpression of GFP-LD (nos-GAL4>UAS-GFP-LD). GFP- and LipidTox Red-positive punctae indicate germline LD (arrows); LipidTox punctae without GFP indicate non-germline LD (arrowheads). Scale bars=10 μm. (H) Histogram representing the spatial distribution of LD within the germline; error bars represent SEM. (I) Histogram representing the spatial distribution of LD and GFP fluorescence (green line) (arbitrary units, a.u.) in a representative testis of a bam-GFP animal (panel J). (J) Testis LD in a bam-GFP animal; arrows point to LD and arrowheads point to spermatocytes. Scale bar=50 μm. See also Supplemental Figure 1.

Drosophila spermatogenesis requires the co-development and differentiation of the germline and the somatic lineages [57]. To identify LD in each lineage, we used the GAL4/UAS system to overexpress a GFP transgene fused to the LD-targeting motif of motor protein Klarsicht (UAS-GFP-LD) [58]. Somatic overexpression of UAS-GFP-LD using Traffic jam (Tj)-GAL4 revealed that the majority of somatic LD in 0-day-old males were located <30 μm from the hub (Figure 1D, 1E). Because the somatic LD distribution coincided with a marker for somatic stem cells and their immediate daughter cells (Zinc finger homeodomain 1, Zfh-1) (Figure 1F; two-sample Kolmogorov-Smirnov test) [59], but not with a marker for late somatic cells (Eyes absent, Eya) [12,60], our data suggest LD are present in early somatic cells. Germline overexpression of UAS-GFP-LD using nanos (nos)-GAL4 demonstrated the presence of LD within germline cells near the apical tip of the testis in 0-day-old males (Figure 1G, 1H). Specifically, the disappearance of germline LD coincided with peak expression of a GFP reporter that reflects the expression of Bag-of-marbles (Bam) protein in the testis (Bam-GFP) [61] (Figure 1I, 1J). Because peak Bam expression signals the last round of transient amplifying mitotic cell cycle prior to the germline’s transition into the meiotic cell cycle [62–64], our data suggests that germline LD, like somatic LD, are present in cells at early stages of development.

brummer plays a cell-autonomous role in regulating testis lipid droplets

Earlier version: Adipose triglyceride lipase (ATGL) is a critical regulator of neutral lipid metabolism and LD [65–74]. Loss of ATGL in many cell types triggers LD accumulation, and ATGL overexpression decreases LD number [30,67,68,71,73,75,76]. Given that the Drosophila ATGL homolog brummer (bmm) regulates testis LD induced by mitochondrial stress [49], we explored whether bmm regulates testis LD under normal physiological conditions. We first examined bmm expression in the testis by isolating this organ from flies carrying a bmm promoter-driven GFP transgene (bmm-GFP) that recapitulates many aspects of bmm mRNA regulation [77]. GFP expression was present in the germline of bmm-GFP testes, and we found germline GFP levels were higher in spermatocytes than at earlier stages of sperm development (Figure 2A, 2B; one-way ANOVA with Tukey multiple comparison test). In further support of this observation, we analyzed a publicly available single-cell RNA sequencing dataset from the male reproductive organ [78]. Using pseudotime analysis, we arranged the germline (Figure S2A) and somatic cells (Figure S2B) based on their annotated developmental trajectory. The expression pattern of bmm in the germline matched our observation with the bmm-GFP reporter (Figure S2C). While levels of the bmm-GFP reporter were low in somatic cells, single-cell RNA sequencing data identified bmm expression in the somatic lineage, with higher expression in cells at later stages of development (Figure S2D).

Revised version: Adipose triglyceride lipase (ATGL) is a critical regulator of neutral lipid metabolism and LD [65–74]. Loss of ATGL in many cell types triggers LD accumulation, and ATGL overexpression decreases LD number [30,67,68,71,73,75,76]. Given that the Drosophila ATGL homolog brummer (bmm) regulates testis LD induced by mitochondrial stress [49], we explored whether bmm regulates testis LD under normal physiological conditions. We first examined bmm expression in the testis by isolating this organ from flies carrying a bmm promoter-driven GFP transgene (bmm-GFP) that recapitulates many aspects of bmm mRNA regulation [77]. GFP expression was present in the germline of bmm-GFP testes, and we found germline GFP levels were higher in spermatocytes than at earlier stages of sperm development (Figure 2A, 2B; one-way ANOVA with Tukey multiple comparison test). In further support of this observation, we analyzed a publicly available single-cell RNA sequencing dataset from the male reproductive organ [78]. Using pseudotime analysis, we arranged the germline (Figure S2A) and somatic cells (Figure S2B) based on their annotated developmental trajectory. The expression pattern of bmm in the germline matched our observation with the bmm-GFP reporter (Figure S2C). While levels of the bmm-GFP reporter were lower in somatic cells, single-cell RNA sequencing data identified bmm expression in the somatic lineage that was higher in cells at later stages of development (Figure S2D). Additional neutral lipid- and lipid droplet-associated genes such as lipid storage droplet-2, Seipin, Lipin, and midway also showed differential regulation during differentiation in the testis (Figure S2C, S2D). Combined with our data on the location of testis LD, these gene expression data suggest that bmm upregulation in both somatic and germline cells during differentiation corresponds to the downregulation of testis LD. Supporting this, germline GFP levels were negatively correlated with testis LD in bmm-GFP flies (Figure 2A, 2C), suggesting regions with higher bmm expression had fewer LD.

brummer regulates lipid droplets in both germline and somatic cells of the testis.

(A-A’’’’) Testis lipid droplets (LD) indicated by LipidTox Red in bmm-GFP animals. Arrows point to LD in all images. Arrowheads point to spermatocytes. Scale bars=50 μm. Asterisks indicate the hub in all images. (B) Quantification of nuclear GFP intensity in testes isolated from bmm-GFP animals (n=3). Germline stem cell (GSC), spermatogonia (SG), spermatocyte (SC). (C) Spatial distribution of LD (grey histogram) and GFP expression (green line) in testes from bmm-GFP animals as a function of distance from the hub (n=3). (D,E) LD near the apical region of the testis in bmmrev (D) or bmm1 (E) animals. (F) LD further away from the apical tip in bmm1 animals. (D–F) Scale bars=50 μm. (G) Histogram representing testis LD size distribution in bmmrev (grey) and bmm1 (orange). (H) Apical tip of the testes is at the left of the graph; individual dots represent a single LD and its relative position to the hub marked by an asterisk. Cumulative frequency distribution of the distance between LDs and the apical tip of the testes are drawn as solid lines. See also Supplemental Figure 2 and 3.

Removed in the revised version: Additional neutral lipid- and lipid droplet-associated genes such as lipid storage droplet-2, Seipin, Lipin, and midway also showed differential regulation during differentiation in the testis (Figure S2C, S2D). Combined with our data on the location of testis LD, these gene expression data suggest that bmm upregulation in both somatic and germline cells during differentiation corresponds to the downregulation of testis LD. Supporting this, germline GFP levels were negatively correlated with testis LD in bmm-GFP flies (Figure 2A, 2C), suggesting regions with higher bmm expression had fewer LD.

Earlier version: To test whether bmm regulates testis LD, we compared LD in testes from 0-day-old males carrying a loss-of-function mutation in bmm (bmm1) to control male testes (bmmrev) [67]. bmm1 males had significantly more LD across all LD sizes compared with control males at the apical tip of the testis (Figure 2D–2G; Welch two-sample t-test with Bonferroni correction) and showed a significantly expanded LD distribution (Figure 2D– 2F, 2H; two-sample Kolmogorov-Smirnov test). This suggests bmm normally restricts LD to the apical tip of the testis, an observation we confirmed in both somatic and germline lineages using lineage-specific expression of GFP-LD (Figure S3A–S3D). Importantly, after inducing homozygous bmmrevor bmm1 clones in the testes using the FLP-FRT system (Figure 3A, 3B) [79], we found bmm1 spermatocyte clones had significantly more LD at 3 days post-clone induction (Figure 3C; Welch two-sample t-test), a stage at which LD were absent from bmmrev clones. Because we observed no significant effect of cell-autonomous bmm loss on LD at any other stage of germline development (Figure 3C), this suggests bmm does not regulate LD at early stages of germ cell development. Instead, our data suggest bmm plays a role in regulating LD at the spermatogonia-spermatocyte transition. While we were unable to assess LD in bmm1 somatic clones, our data reveals a previously unrecognized cell-autonomous role for bmm as a key regulator of LD in germline cells.

Revised version: To test whether bmm regulates testis LD, we compared LD in testes from 0-day- old males carrying a loss-of-function mutation in bmm (bmm1) to control male testes (bmmrev) [67]. bmm1 males had significantly more LD across all LD sizes compared with control males at the apical tip of the testis (Figure 2D–2G; Welch two-sample t-test with Bonferroni correction) and showed a significantly expanded LD distribution (Figure 2D– 2F, 2H; two-sample Kolmogorov-Smirnov test). This suggests bmm normally restricts LD to the apical tip of the testis, an observation we confirmed in both somatic and germline lineages using lineage-specific expression of GFP-LD (Figure S3A–S3D). Importantly, after inducing homozygous bmmrev or bmm1 clones in the testes using the FLP-FRT system (Figure 3A, 3B) [79], we found bmm1 spermatocyte clones had significantly more LD at 3 days post-clone induction (Figure 3C; Welch two-sample t-test), a stage at which LD were absent from bmmrev clones. Because we observed no significant effect of cell-autonomous bmm loss on LD at any other stage of germline development (Figure 3C), this suggests bmm function is not required to regulate LD at early stages of germ cell development. Instead, our data suggest bmm plays a role in regulating LD at the spermatogonia-spermatocyte transition. While we were unable to assess LD in bmm1 somatic clones, our data reveals a previously unrecognized cell-autonomous role for bmm as a regulator of LD in germline cells.

bmm regulates germline lipid droplets in a cell-autonomous manner.

Earlier version: (A–B) Single confocal slices through a representative testis isolated from an individual carrying clones induced using the FLP-FRT system at 3 days post-clone induction. Clones are homozygous for an allele that encodes a functional bmm protein product (bmmrev; A–A’’’) or a loss-of-function bmm allele (bmm1, B–B’’’). GFP negative areas mark homozygous clones in panel A and B; the boxed areas in A, A’’ and B, B’’ are shown in A’, A’’’ and B’, B’’’, respectively. In homozygous bmmrev spermatocyte clones we detected no LD using neutral lipid dye BODIPY (A’’ and A’’’). In contrast, spermatocyte clones homozygous for bmm1 have detectable LD (B’’ and B’’’, arrows). Scale bars=50 μm in A,A’’ and B,B’’; scale bars=10 μm in A’,A’’’ and B’,B’’’. (C) Number of testis LD in bmmrev (grey) or bmm1 (orange) in FLP-FRT clones 3 days post-clone induction; dots represent measurements from a single clone. The number of cells in each cyst (CC) counted is indicated. There were significantly more LD in bmm1 spermatocyte (SC) clones (p = 0.026; Welch two-sample t-test) but not at other stages of development.

Revised version: (A–B) Single confocal slices through a representative testis isolated from an individual carrying clones induced using the FLP-FRT system at 3 days post-clone induction. Clones are homozygous for an allele that encodes a functional bmm protein product (bmmrev; A–A’’’) or a loss-of-function bmm allele (bmm1, B–B’’’). GFP negative areas mark homozygous clones in panel A and B; the boxed areas in A, A’’ and B, B’’ are shown in A’, A’’’ and B’, B’’’, respectively. In homozygous bmmrev spermatocyte clones we detected no LD using neutral lipid dye LipidTox (A’’ and A’’’). In contrast, spermatocyte clones homozygous for bmm1 have detectable LD (B’’ and B’’’, arrows). Scale bars=50 μm in A,A’’ and B,B’’; scale bars=10 μm in A’,A’’’ and B’,B’’’. (C) Number of testis LD in bmmrev (grey) or bmm1 (orange) in FLP-FRT clones 3 days post-clone induction; dots represent measurements from a single clone. The number of cells in each cyst (CC) counted is indicated. There were significantly more LD in bmm1 spermatocyte (SC) clones (p = 0.026; Welch two-sample t-test) but not at other stages of development.

brummer plays a cell-autonomous role in regulating germline development

Earlier version: To determine the physiological significance of bmm-mediated regulation of testis LD, we investigated testis and sperm development in males without bmm function. In 0-day-old bmm1 males reared at 25°C, testis size was significantly smaller than in age-matched bmmrev controls (Figure 4A, 4B; Welch two-sample t-test), and the number of spermatid bundles was significantly lower (Figure 4C; Kruskal-Wallis rank sum test). When the animals were reared at 29°C, a temperature that exacerbates spermatogenesis defects associated with changes in lipid metabolism [21], bmm1 phenotypes were more pronounced (Figure S3A, S3B; Welch two-sample t-test, Kruskal-Wallis rank sum test). Defects in testis size were also observed at 14 days post-eclosion; suggesting testis size defects persist later into the life course (Figure S4C; Welch two-sample t-test). In contrast, the number of spermatid bundles per testis was not significantly different between bmm1 and bmmrevmales at this age (Figure S4D; Welch two-sample t-test), potentially due to a large decrease in the number of spermatid bundles in 14-day-old bmmrev males (Figure 4C, S4D). Together, these data suggest loss of bmm affects testis development and spermatogenesis. Similar phenotypes are observed in male mice without ATGL [52], and supplementing the diet of bmm1males with medium-chain triglycerides (MCT) partially rescues the testis and spermatogenic defects we observed in flies (Figure S4E, S4F; one-way ANOVA with Tukey multiple comparison test), as it does in mice [52,80]. This identifies similarities between flies and mice in fertility-related phenotypes associated with whole-body loss of bmm/ATGL.

Revised version: To determine the physiological significance of bmm-mediated regulation of testis LD, we investigated testis and sperm development in males without bmm function. In 0-day-old bmm1 males reared at 25°C, testis size was significantly smaller than in age-matched bmmrev controls (Figure 4A, 4B; Welch two-sample t-test), and the number of spermatid bundles was significantly lower (Figure 4C; Kruskal-Wallis rank sum test). When the animals were reared at 29°C, a temperature that exacerbates spermatogenesis defects associated with changes in lipid metabolism [21], bmm1 phenotypes were more pronounced (Figure S4A, S4B; Welch two-sample t-test, Kruskal-Wallis rank sum test). Defects in testis size were also observed at 14 days post-eclosion; suggesting testis size defects persist later into the life course (Figure S4C; Welch two-sample t-test). In contrast, the number of spermatid bundles per testis was not significantly different between bmm1 and bmmrev males at this age (Figure S4D; Welch two-sample t-test), potentially due to a large decrease in the number of spermatid bundles in 14-day-old bmmrev males (Figure 4C, S4D). Together, these data suggest loss of bmm affects testis development and spermatogenesis. Similar phenotypes are observed in male mice without ATGL [52], and supplementing the diet of bmm1 males with medium-chain triglycerides (MCT) partially rescues the testis and spermatogenic defects we observed in flies (Figure S4E, S4F; one-way ANOVA with Tukey multiple comparison test), as it does in mice [52,80]. This identifies similarities between flies and mice in fertility-related phenotypes associated with whole-body loss of bmm/ATGL.

A cell-autonomous role for bmm in regulating spermatogenesis.

Testes isolated from bmmrev (A) and bmm1(A’) animals raised at 25°C stained with phalloidin. Scale bars=100 μm. (B) Testis size in bmm1 and bmmrev animals raised at 25°C. (C) Spermatid bundle number in bmm1 and bmmrev testes from animals reared at 25°C. (D,E) Representative images of bmmrev (D) or bmm1 (E) testes stained with DAPI and anti-Vasa antibody. Arrows indicate germline stem cells (GSC). Scale bar=50 μm. The hub is marked by an asterisk in all images. (F) GSC number in bmm1 and bmmrev testes. (G) Proportion of GSCs that were either bmm1 or bmmrevclones at 3 and 14 days post-clone induction. (H) Representative images of bmmrev(H) and bmm1 (H’) testes carrying bam-GFP; data quantified in Figure S4J. Arrows indicate regions with high Bam-GFP. Scale bars=50 μm. (I) Representative images of bmmrev (I) or bmm1(I’,I’’) testes stained with anti-Vasa antibody. Arrows indicate Vasa-positive cysts in bmm1 testis. Panel I’’ is magnified from the boxed region in I’. (I,I’) Scale bars=100 μm; (I’’) scale bar=50 μm. (J) Maximum projection of bmmrev (J) or bmm1 (J’) testes stained with anti-Boule antibody (green) and DAPI (blue). Scale bars=100 μm. Number of bmm1and bmmrev spermatocyte clones (K) or post-meiotic clones (L) at 3 and 14 days post-clone induction. See also Supplemental Figure 4.

Earlier version: To explore spermatogenesis in bmm1 animals, we used an antibody against germ cell-specific marker Vasa to visualize the germline in the testes of bmm1 and bmmrevmales (Figure 4D, 4E) [81]. We observed a significant increase in the number of germline stem cells (GSC) (Figure 4F; Kruskal-Wallis rank sum test) and higher variability in GSC number in bmm1 males (p=5.7×10-12 by F-test). Given that GSC number is affected by hub size and GSC proliferation [82,83], we monitored both parameters in bmm1 and bmmrev controls. While hub size in bmm1 testes was significantly larger than in testes from bmmrev controls (Figure S4G, S4H; Welch two-sample t-test), the number of phosphohistone H3-positive GSCs, which indicates proliferating GSCs, was unchanged in bmm1 animals (Figure S4I; Kruskal-Wallis rank sum test). While this indicates a larger hub may partly explain bmm’s effect on GSC number, bmm also plays a cell-autonomous role in regulating GSCs, as we recovered a higher proportion of bmm1 clones in the GSC pool compared with bmmrevclones at 14 days after clone induction (Figure 4G; Welch two-sample t-test). Given that we detected no effect of cell-autonomous bmm loss on the number of GSC LD (Figure 3C), more work will be needed to understand how bmm regulates GSCs at a stage prior to its effects on LD number. Future studies will also need to confirm whether bmm1 mutant GSCs show an increased ability to occupy space at the hub.

Revised version: To explore spermatogenesis in bmm1 animals, we used an antibody against the germline cell-specific marker Vasa to visualize the germline in the testes of bmm1 and bmmrev males (Figure 4D, 4E) [81]. We observed a significant increase in the number of germline stem cells (GSC) (Figure 4F; Kruskal-Wallis rank sum test) and higher variability in GSC number in bmm1 males (p=5.7×10-12 by F-test). Given that GSC number is affected by hub size and GSC proliferation [82,83], we monitored both parameters in bmm1 and bmmrev controls. While hub size in bmm1 testes was significantly larger than in testes from bmmrev controls (Figure S4G, S4H; Welch two-sample t-test), the number of phosphohistone H3-positive GSCs, which indicates proliferating GSCs, was unchanged in bmm1 animals (Figure S4I; Kruskal-Wallis rank sum test). While this indicates a larger hub may partly explain bmm’s effect on GSC number, bmm also plays a cell-autonomous role in regulating GSCs, as we recovered a higher proportion of bmm1 clones in the GSC pool compared with bmmrev clones at 14 days after clone induction (Figure 4G; Welch two-sample t-test). Given that we detected no effect of cell-autonomous bmm loss on the number of GSC LD (Figure 3C), more work will be needed to understand how bmm regulates GSCs at a stage prior to its effects on LD number. Future studies will also need to confirm whether bmm1 mutant GSCs show an increased ability to occupy space at the hub.

Removed in the revised version: Beyond GSCs, we uncovered additional spermatogenesis defects in bmm1testes. Peak Bam-GFP expression in germline cells of the testes from 0-day-old bmm1 and bmmrev males showed GFP-positive cysts were significantly further away from the hub in bmm1testes (Figure 4H, S4J; Welch two-sample t-test). Indeed, 15/18 bmm1 testes contained Vasa-positive cysts with large nuclei in the distal half of the testis (Figure 4I, arrows), a phenotype not present in bmmrev testes (0/8) (p=0.0005 by Pearson’s Chi-square test). Because these phenotypes are also seen in testes with differentiation defects [13,84], we recorded the stage of sperm development reached by the germline in bmm1 testes. Most bmm1 testes contained post-meiotic cells in males raised at 25°C (Figure S4K); however, germline development did not progress past the spermatocyte stage in most bmm1 testes from animals raised at 29°C (Figure S4K).

Earlier version: Testes from bmm1 males reared at 25°C also had a smaller Boule-positive area (Figure 4J, S4L; Welch two-sample t-test), and fewer individualization complexes and waste bags (Figure S4M, S4N; Kruskal-Wallis rank sum test). Because Boule-positive area, individualization complexes, and waste bags indicate germline cells at later stages in development, these data suggest the loss of bmm causes a reduction in differentiated cell types. Because we observed significantly fewer bmm1 spermatocyte and spermatid clones at 14 days after clone induction (Figure 4K, 4L; p = 0.0496, Kruskal-Wallis rank sum test), these effects on germline development may represent a cell-autonomous role for bmm in regulating spermatogenesis in this cell type. Given that the statistical significance of this finding was not as strong as for our other data, future studies should repeat this experiment with more samples. We also reveal a potential non-cell-autonomous role for somatic bmm. While there was no difference in the ratio of Zfh-1-positive cells between homozygous clones and heterozygous clones in animals carrying the bmm1 or bmmrevalleles at 14 days post clone induction (Figure S4O; Kruskal-Wallis rank sum test), the distance from the hub to the Zfh-1 positive clones reside was significantly decreased in bmm1 homozygous clones (Figure S4P; Kruskal-Wallis rank sum test). Together, these data indicate bmm plays a cell-autonomous role in germline cells, and potentially a non-cell-autonomous role in somatic cells, to regulate spermatogenesis.

Revised version: Beyond GSCs, we uncovered additional spermatogenesis defects in bmm1testes. Peak Bam-GFP expression in germline cells of the testes from 0-day-old bmm1 and bmmrev males showed GFP-positive cysts were significantly further away from the hub in bmm1testes (Figure 4H, S4J; Welch two-sample t-test). Indeed, 15/18 bmm1 testes contained Vasa-positive cysts with large nuclei in the distal half of the testis (Figure 4I, arrows), a phenotype not present in bmmrev testes (0/8) (p=0.0005 by Pearson’s Chi-square test). Because these phenotypes are also seen in testes with differentiation defects [13,84], we recorded the stage of sperm development reached by the germline in bmm1 testes. Most bmm1 testes contained post-meiotic cells in males raised at 25°C (Figure S4K); however, germline development did not progress past the spermatocyte stage in most bmm1 testes from animals raised at 29°C (Figure S4K). Testes from bmm1 males reared at 25°C also had a smaller Boule-positive area (Figure 4J, S4L; Welch two-sample t-test), and fewer individualization complexes and waste bags (Figure S4M, S4N; Kruskal-Wallis rank sum test). Because Boule-positive area, individualization complexes, and waste bags are all markers for later stages in sperm development, these data indicate the loss of bmm causes a reduction in differentiated cell types. Because we observed significantly fewer bmm1 spermatocyte and spermatid clones at 14 days after clone induction (Figure 4K, 4L; p = 0.0496, Kruskal-Wallis rank sum test), these effects on germline development may represent a cell-autonomous role for bmm in regulating spermatogenesis in this cell type. Given that the statistical significance of this finding was not as strong as for our other data, future studies should repeat this experiment with more samples. We also reveal a potential non-cell-autonomous role for somatic bmm. While there was no difference in the ratio of Zfh-1-positive cells between homozygous clones and heterozygous clones in animals carrying the bmm1 or bmmrevalleles at 14 days post clone induction (Figure S4O; Kruskal-Wallis rank sum test), the distance from the hub to the Zfh-1 positive clones was significantly decreased in bmm1 homozygous clones (Figure S4P; Kruskal-Wallis rank sum test). Together, these data indicate bmm may play a cell-autonomous role in germline cells, and potentially a non-cell-autonomous role in somatic cells, to regulate spermatogenesis.

brummer-dependent regulation of testis triglyceride levels affects spermatogenesis

ATGL catalyzes the first and rate-limiting step of triglyceride hydrolysis [73,85,86]. Loss of this enzyme or its homologs leads to excess triglyceride accumulation [27,30,67,73,75] and shifts in multiple lipid classes [66,87–89]. To determine how loss of bmm affects spermatogenesis, we carried out whole-body mass spectrometry (MS)- based untargeted lipidomic profiling of bmm1 and bmmrev males. Hierarchical clustering of lipid species suggests that bmm1 and bmmrev males show distinct lipidomic profiles (Figure 5A). Overall, we detected 2464 and 1144 lipid features with high quantitative confidence in positive and negative ion modes, respectively. By matching experimental m/z, isotopic ratio, and tandem MS spectra to lipid libraries, we confirmed 293 unique lipid species (Supplemental table 1). We found 107 lipids had a significant change in abundance between bmm1 and bmmrev males (padj<0.05): 85 species were upregulated in bmm1 males and 22 lipid species were downregulated. Among differentially regulated species from different lipid classes, triglyceride had the largest residual above expected proportion (p=5.00×10-4 by Pearson’s Chi-squared test). This suggests triglyceride is the lipid class most affected by loss of bmm (Figure 5B, 5C).

Loss of bmm disrupts triglyceride homeostasis and leads to spermatogenic defects.

Earlier version: (A) Hierarchical clustering of lipid species detected in bmmrev and bmm1 animals. (B) Histograms showing the proportion of species in each lipid class with different levels between bmm1 and bmmrev. Numbers on histograms indicate the number of species with differences in abundance. (C) Volcano plot showing fold change in abundance of triglycerides (green; 97 species) and non-triglyceride lipids (grey; 186 species) in our dataset. (D) Arrows indicate testis LD stained with LipidTox Red in bmmrev (D), bmm1 (D’), or mdyQX25/k03902; bmm1 (D’’) animals. (E) Whole testes isolated from bmmrev(E), bmm1 (E’), or mdyQX25/k03902;bmm1(E’’) animals stained with anti-Vasa antibody (red) and DAPI (blue). Arrowheads indicate spermatid bundles. Scale bars=100 μm. (F) Testis size in bmmrev, bmm1, and mdyQX25/k03902;bmm1animals. Spermatid bundles (G) and number of germline stem cells (H) in bmmrev, bmm1, and mdyQX25/k03902;bmm1 animals. (I) Testis size in animals with germline-specific mdy knockdown (nos-GAL4>mdy RNAi; bmm1) compared with controls (nos-GAL4>+; bmm1 and +>mdy RNAi; bmm1). See also Supplemental Figure 5.

Revised version: (A) Hierarchical clustering of lipid species detected in bmmrev and bmm1 animals. (B) Histograms showing the proportion of significant species in each lipid class with different levels between bmm1 and bmmrev. Numbers on histograms indicate the number of species with differences in abundance. (C) Volcano plot showing fold change in abundance of triglycerides (green; 97 species) and non-triglyceride lipids (grey; 186 species) in our dataset. (D) Arrows indicate testis LD stained with LipidTox Red in bmmrev (D), bmm1 (D’), or mdyQX25/k03902; bmm1 (D’’) animals. (E) Whole testes isolated from bmmrev(E), bmm1 (E’), or mdyQX25/k03902;bmm1(E’’) animals stained with anti-Vasa antibody (red) and DAPI (blue). Arrowheads indicate spermatid bundles. Scale bars=100 μm. (F) Testis size in bmmrev, bmm1, and mdyQX25/k03902;bmm1 animals. Spermatid bundles (G) and number of germline stem cells (H) in bmmrev, bmm1, and mdyQX25/k03902;bmm1 animals. (I) Testis size in animals with germline-specific mdy knockdown (nos-GAL4>mdy RNAi; bmm1) compared with controls (nos-GAL4>+; bmm1 and +>mdy RNAi; bmm1). See also Supplemental Figure 5.

In bmm1 males, the majority of triglyceride species (55/97) were significantly higher in abundance compared with bmmrevcontrol males. Because we observed a positive correlation between the fold increase in triglyceride abundance with both the number of double bonds (p=7.52×10-8 by Kendall’s rank correlation test; Figure S5A) and the number of carbons (p=2.77×10-10 by Kendall’s rank correlation test; Figure S5B), our data align well with bmm/ATGL’s known role in regulating triglyceride levels [67,68,73] and its substrate preference of long-chain polyunsaturated fatty acids [85]. While we also detected changes in species such as fatty acids, acylcarnitine, and membrane lipids (Figure S5C–S5H), in line with recent Drosophila lipidomic data [90,91], the striking accumulation of triglyceride in bmm1 males suggested that excess testis triglyceride in bmm1 males may contribute to their spermatogenic defects. To test this, we examined spermatogenesis in bmm1 males carrying loss-of-function mutations in midway (mdy). mdy is the Drosophila homolog of diacylglycerol O-acyltransferase 1 (DGAT1), and whole-body loss of mdy reduces whole-body triglyceride levels [92–94]. Importantly, testes isolated from males with global loss of both bmm and mdy (mdyQX25/k03902;bmm1) had fewer LD than testes dissected from bmm1 males (Figures 5D, S5I; one-way ANOVA with Tukey multiple comparison test).

We found that testes isolated from mdyQX25/k03902;bmm1males were significantly larger and had more spermatid bundles than testes from bmm1 males (Figure 5E–G; one-way ANOVA with Tukey multiple comparison test). The elevated number of GSCs in bmm1 male testes was similarly rescued in mdyQX25/k03902;bmm1 males (Figure 5H; one-way ANOVA with Tukey multiple comparison test). These data suggest that defective spermatogenesis in bmm1 males can be partly attributed to excess triglyceride accumulation. Notably, at least some of the effects of global mdy loss on bmm1 males can be attributed to the germline: RNAi-mediated knockdown of mdy in the germline of bmm1 males partially rescued the defects in testis size (Figure 5I; Kruskal-Wallis rank sum test with Dunn’s multiple comparison test) and GSC variance (Figure S5J; p=4.5 x 10-5 and 8.2 x 10-3 by F-test from the GAL4- and UAS-only crosses, respectively). While future studies will need to test whether germline-specific loss of mdy also rescues spermatid number defects in bmm1 males, our data suggest bmm-mediated regulation of testis triglyceride plays a previously unrecognized role in regulating sperm development.

Discussion

Earlier version: In this study, we used Drosophila to gain insight into how the neutral lipids, a major lipid class, contribute to sperm development. We describe the distribution of LD under normal physiological conditions in the Drosophila testis, and show that LD are present at the early stages of development in both somatic and germline cells. While many factors are known to regulate LD in nongonadal cell types, we reveal a cell-autonomous role for triglyceride lipase bmm in regulating testis LD during spermatogenesis. In particular, we identified a requirement for bmm in mediating the decrease in LD at the spermatogonia-spermatocyte transition. This regulation is important for sperm development, as our data indicates that loss of bmm causes a decrease in the number of differentiated cell types in the testis. This phenotype may be attributed to a delay in differentiation, a block in differentiation, or to a loss of differentiated cells through cell death. Future studies will therefore be essential to resolve why bmm loss causes a reduction in differentiated cell types. Nevertheless, these defects in the number of differentiated cell types can be partially explained by the excess accumulation of triglyceride in flies lacking bmm, as global and cell type-specific inhibition of triglyceride synthesis rescues multiple spermatogenic defects in bmm mutants. Together, our data reveals previously unrecognized roles for LD and triglycerides during spermatogenesis, and for bmm as an important regulator of testis LD and germline development under normal physiological conditions.

Revised version: In this study, we used Drosophila to gain insight into how the neutral lipids, a major lipid class, contribute to sperm development. We describe the distribution of LD under normal physiological conditions in the Drosophila testis, and show that LD are present at the early stages of development in both somatic and germline cells. While many factors are known to regulate LD in nongonadal cell types, we reveal a cell-autonomous role for triglyceride lipase bmm in regulating testis LD during spermatogenesis. In particular, we identified a requirement for bmm in mediating the decrease in LD at the spermatogonia-spermatocyte transition. This regulation is important for sperm development, as our data indicates that loss of bmm causes a decrease in the number of differentiated cell types in the testis. This reduction in differentiated cell types may be attributed to a delay in differentiation, a block in differentiation, or to a loss of differentiated cells through cell death. Future studies will therefore be essential to resolve why bmm loss causes a reduction in differentiated cell types. Nevertheless, these defects in the number of differentiated cell types can be partially explained by the excess accumulation of triglyceride in flies lacking bmm, as global and cell type-specific inhibition of triglyceride synthesis rescues multiple spermatogenic defects in bmm mutants. Together, our data reveals previously unrecognized roles for LD and triglycerides during spermatogenesis, and for bmm as an important regulator of testis LD and germline development under normal physiological conditions.

One key outcome of our study was increased knowledge of LD regulation and function in the testis. Despite rapidly expanding knowledge of LD in cell types such as adipocytes or skeletal muscle, less is known about how LD influence spermatogenesis under normal physiological conditions. In mammals, testis LD contain cholesterol and play a role in promoting steroidogenesis [95,96]. In flies, we show that LD are present in the testis, and that excess accumulation of these LD affects sperm development. In nongonadal cell types, triglycerides provide a rich source of fatty acids for cellular ATP production, lipid building blocks to support membrane homeostasis and growth, and metabolites that can act as signaling molecules [26]. Because ATP production, lipid precursors, and lipid signaling all play roles in supporting normal sperm development [97,98], future studies will need to determine how each of these processes is affected when excess triglyceride accumulates in testis LD (Figure 6). It will also be important to determine whether it is the loss of metabolites produced by bmm’s enzymatic action, or an increase in triglycerides, that leads to the reduction in differentiated cell types during spermatogenesis. Together, these experiments will provide critical insight into how triglyceride stored within testis LD contributes to overall cellular lipid metabolism during spermatogenesis. Because of the parallel spermatogenic defects we observed in bmm mutants and ATGL-deficient mice, we expect these mechanisms will also operate in other species.

Model of bmm-mediated lipid droplet regulation in the Drosophila testis.

Schematic representation summarizing bmm-mediated lipid droplet regulation in the testis during development.

Removed in the revised version: A more comprehensive understanding of neutral lipid metabolism during sperm development will also emerge from studies on the upstream signaling networks that regulate testis LD and triglyceride. Given that we show an important and cell-autonomous role for bmm in regulating testis LD and triglyceride, future studies will need to identify factors that regulate bmm in the testis. Based on public single-cell RNAseq data and the bmm-GFP reporter strain, our data suggest bmm mRNA levels are differentially regulated between early and later stages of sperm development.

Removed in the revised version: Candidates for mediating this regulation include the insulin/insulin-like growth factor signaling pathway (IIS), Target of rapamycin (TOR) pathway, and nuclear factor κB/Relish pathway (NFκB), as all of these pathways influence bmm mRNA levels in nongonadal cell types [99–105]. Beyond mRNA levels, Bmm protein levels and post-translational modifications may also be differentially regulating during spermatogenesis.

Earlier version: For example, studies show that the proteins encoded by bmm homologs in other animals are regulated by phosphorylation [106], mediated by kinases such as adenosine monophosphate-activated protein kinase (AMPK) and protein kinase A (PKA) [107–109]. Importantly, many of these pathways, including IIS, TOR, AMPK, NFκB and possibly PKA influence Drosophila sperm development [110–115]. Identifying the signaling networks that influence bmm regulation during sperm development will therefore lead to a deeper understanding of how testis LD and triglyceride are coordinated with physiological factors to promote normal spermatogenesis. Because pathways such as IIS and AMPK, and others, regulate sperm development in other species [116–118], these insights may reveal conserved mechanisms that govern the regulation of cellular neutral lipid metabolism during sperm development.

Revised version: A more comprehensive understanding of neutral lipid metabolism during sperm development will also emerge from studies on the upstream signaling networks that regulate testis LD and triglyceride. Given that we show an important and cell-autonomous role for bmm in regulating testis LD and triglyceride, future studies will need to identify factors that regulate bmm in the testis. Based on public single-cell RNAseq data and the bmm-GFP reporter strain, our data suggest bmm mRNA levels are differentially regulated between early and later stages of sperm development. Candidates for mediating this regulation include the insulin/insulin-like growth factor signaling pathway (IIS), Target of rapamycin (TOR) pathway, and nuclear factor κB/Relish pathway (NFκB), as all of these pathways influence bmm mRNA levels in nongonadal cell types [99–105]. Beyond mRNA levels, Bmm protein levels and post- translational modifications may also be differentially regulating during spermatogenesis. For example, studies show that the proteins encoded by bmm homologs in other animals are regulated by phosphorylation [106], mediated by kinases such as adenosine monophosphate-activated protein kinase (AMPK) and protein kinase A (PKA) [107–109]. Importantly, many of these pathways, including IIS, TOR, AMPK, NFκB and possibly PKA influence Drosophila sperm development [110–115]. Identifying the signaling networks that influence bmm regulation during sperm development will therefore lead to a deeper understanding of how testis LD and triglyceride are coordinated with physiological factors to promote normal spermatogenesis. Because pathways such as IIS and AMPK, and others, regulate sperm development in other species [116–118], these insights may reveal conserved mechanisms that govern the regulation of cellular neutral lipid metabolism during sperm development.

Supporting information